### Lecture 3-2 Practice: Lambda Functions and Dictionaries

### Question 1
Sorting with lambda key functions.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
# Sort list of numbers by absolute value
nums = [-4, 1, -2, 3]
nums.sort(key=lambda x: abs(x))
assert(nums == [1, -2, 3, -4])
print(nums)

# Sort list of strings by length.
words = ['apple', 'cat', 'banana'] 
words.sort(key=lambda x: len(x))
assert(words == ['cat', 'apple', 'banana'])
print(words)
```

Output:
```
[1, -2, 3, -4]
['cat', 'apple', 'banana']
```

### Question 2
Given an array of intervals where `intervals[i] = [starti, endi]`, merge all overlapping intervals, and return an array of the non-overlapping intervals that cover all the intervals in the input. (Note: this is a classical interview question from Leetcode)

Example:

Input: [[1,3],[2,6],[8,10],[15,18]]

Output: [[1,6],[8,10],[15,18]]

```python
# def merge_intervals(intervals):
#    y = []
#    z = []
#    for x in intervals:
#      if x[0] < y[1]:
#            z.pop()
#            z.append([y[0],x[1]])
#        else:
#            z.append(x)
#        y = x
#    return z


```

```python
assert merge_intervals([[1,3],[2,6],[8,10],[15,18]]) == [[1,6],[8,10],[15,18]]
assert merge_intervals([[1,6],[2,4],[3,5]]) == [[1,6]]
assert merge_intervals([[-4,-1],[-2,3],[2,6]]) == [[-4,6]] 
```

Output:
```
Cell In[22], line 1
----> 1 assert merge_intervals([[1,3],[2,6],[8,10],[15,18]]) == [[1,6],[8,10],[15,18]]
      2 assert merge_intervals([[1,6],[2,4],[3,5]]) == [[1,6]]
      3 assert merge_intervals([[-4,-1],[-2,3],[2,6]]) == [[-4,6]]

Cell In[20], line 5, in merge_intervals(intervals)
      3 z = []
      4 for x in intervals:
----> 5     if x[0] < y[1]:
      6         z.pop()
      7         z.append([y[0],x[1]])

IndexError: list index out of range
```

### Question 3
Define a function called `switch_key_value`, which takes a dictionary as its only argument. Switch the placement of key and value of the dictionary and return it as a new dictionary.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
def switch_key_value(d):
    ans = {}
    for k,v in d.items():
        print("key",k)
        print("value",v)
        ans[v] = k
        print(ans)
    return ans
```

<!-- grade_id: cell-a0b5ecba31ff8afb, points: 1, grade: true, solution: false -->
```python
assert(switch_key_value({}) == {})
assert(switch_key_value({'a': 1, 'b': 2}) == {1: 'a', 2: 'b'})
```

Output:
```
key a
value 1
{1: 'a'}
key b
value 2
{1: 'a', 2: 'b'}
```

### Question 4
Define a function called `sort_list_dict`, which takes a dictionary whose values are lists as its only argument. Sort all these lists alphabetically return them in a new dictionary (with the same keys).

<!-- grade_id: cell-996c61e79b68ca6d, points: 0, grade: false, solution: true -->
```python
input_dict = {'a': [4, 1, 2, 3], 'b': [3, 2, 1]}
def sort_list_dict(d):
    ans={} 
    for k,v in d.items():
         print(v)
         ans[k]=sorted(v)
    return ans
sort_list_dict(input_dict)
```

Output:
```
[4, 1, 2, 3]
[3, 2, 1]
{'a': [1, 2, 3, 4], 'b': [1, 2, 3]}
```

<!-- grade_id: cell-f752b519fbf2ac78, points: 1, grade: true, solution: false -->
```python
assert(sort_list_dict({}) == {})

assert(sort_list_dict({'a': []}) == {'a': []})

input_dict = {'a': [4, 1, 2, 3], 'b': [3, 2, 1]}
desired_output = {'a': [1, 2, 3, 4], 'b': [1, 2, 3]}
assert(sort_list_dict(input_dict) == desired_output)

input_dict = {'S': [2, 3, 1], 'T': [5, 1, 2], 'D': [3, 2, 4]}
desired_output = {'S': [1, 2, 3], 'T': [1, 2, 5], 'D': [2, 3, 4]}
assert(sort_list_dict(input_dict) == desired_output)
```

Output:
```
[]
[4, 1, 2, 3]
[3, 2, 1]
[2, 3, 1]
[5, 1, 2]
[3, 2, 4]
```

```python

```